# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


```python
import pandas as pd
import os
```

```python
os.getcwd()
```

Output:
```
'/home/jovyan/work/assignments/assignment2'
```

```python
# help(os)
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    import pandas as pd
    data=pd.read_csv('assets/NISPUF17.csv',index_col=0)
    data_ed=data.groupby(by=['EDUC1']).count().reset_index()
    data_ed=data_ed[['EDUC1','SEQNUMHH']]
    
    dictionary={"less than high school":data_ed.iloc[0,1]/data_ed.iloc[:,1].sum(),
    "high school":data_ed.iloc[1,1]/data_ed.iloc[:,1].sum(),
    "more than high school but not college":data_ed.iloc[2,1]/data_ed.iloc[:,1].sum(),
    "college":data_ed.iloc[3,1]/data_ed.iloc[:,1].sum()}
    
    return dictionary
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

```python
type(proportion_of_education())==type({})
```

Output:
```
True
```

```python
data=pd.read_csv('/home/jovyan/work/assignments/assignment2/assets/NISPUF17.csv',index_col=0)

```

```python
pd.options.display.max_rows=1000
pd.options.display.max_columns=1000
```

```python
data.head()
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  AGECPOXR  \
1   128521     12852     2         NaN  235.916956     1031  2017       NaN   
2    10741      1074     2         NaN  957.353840     1068  2017       NaN   
3   220011     22001     2         NaN  189.611299     1050  2017       NaN   
4    86131      8613     1  675.430817  333.447418     1040  2017       NaN   
5   227141     22714     1  482.617748  278.768063     1008  2017       NaN   

   HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXCLR06  BF_FORMR08  BFENDFL06  BFFORMFL06  \
1         2       1   395.6875      121.75     182.625        NaN         NaN   
2         2       1        NaN         NaN         NaN        NaN         NaN   
3         2       3        NaN         NaN         NaN        NaN         NaN   
4         2       1        NaN         NaN         NaN        NaN         NaN   
5         2       1   547.8750      273.90     888.000        NaN         NaN   

   C1R  C5R  CBF_01  CEN_REG  CHILDNM  CW
```

```python
data.head()
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  AGECPOXR  \
1   128521     12852     2         NaN  235.916956     1031  2017       NaN   
2    10741      1074     2         NaN  957.353840     1068  2017       NaN   
3   220011     22001     2         NaN  189.611299     1050  2017       NaN   
4    86131      8613     1  675.430817  333.447418     1040  2017       NaN   
5   227141     22714     1  482.617748  278.768063     1008  2017       NaN   

   HAD_CPOX  AGEGRP  BF_ENDR06  BF_EXCLR06  BF_FORMR08  BFENDFL06  BFFORMFL06  \
1         2       1   395.6875      121.75     182.625        NaN         NaN   
2         2       1        NaN         NaN         NaN        NaN         NaN   
3         2       3        NaN         NaN         NaN        NaN         NaN   
4         2       1        NaN         NaN         NaN        NaN         NaN   
5         2       1   547.8750      273.90     888.000        NaN         NaN   

   C1R  C5R  CBF_01  CEN_REG  CHILDNM  CW
```

```python
data['CBF_01'].unique()
```

Output:
```
array([ 1,  2, 99, 77])
```

```python
{'male':0.1, 'female':0.2}   
```

Output:
```
{'male': 0.1, 'female': 0.2}
```

```python
import numpy as np
```

```python
data[(data.CBF_01==1) | (data.CBF_01==2)].groupby('CBF_01')['P_NUMFLU'].agg(np.mean)
```

Output:
```
CBF_01
1    1.879919
2    1.596395
Name: P_NUMFLU, dtype: float64
```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    import pandas as pd
    data=pd.read_csv('assets/NISPUF17.csv',index_col=0)
    a,b=data[(data.CBF_01==1) | (data.CBF_01==2)].groupby('CBF_01')['P_NUMFLU'].agg(np.mean)
    
    return (a,b)
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

```python
data['SEX']
```

Output:
```
1    1
2    1
3    2
4    2
5    2
Name: SEX, dtype: int64
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
data[(data['HAD_CPOX']==2) & (data['P_NUMVRC']>0)].groupby('SEX')['HAD_CPOX'].count()   

```

Output:
```
SEX
1    7028
2    6802
Name: HAD_CPOX, dtype: int64
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    import pandas as pd
    data=pd.read_csv('assets/NISPUF17.csv',index_col=0)
    data=data[['HAD_CPOX','P_NUMVRC','SEX']]
    a,b=data[(data['HAD_CPOX']==1) & (data['P_NUMVRC']>0)].groupby('SEX')['HAD_CPOX'].count()   
    c,d=data[(data['HAD_CPOX']==2) & (data['P_NUMVRC']>0)].groupby('SEX')['HAD_CPOX'].count()   
#     print(a,b,c,d) 
    return {'male':a/c, 'female':b/d}    
    
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

```python
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more nos) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isnt really the full picture, since we are not looking at when the dose was given. Its possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe

    data=pd.read_csv('assets/NISPUF17.csv',index_col=0)

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(data[(~data.P_NUMVRC.isnull()) & (data.HAD_CPOX<3)]["HAD_CPOX"],data[(~data.P_NUMVRC.isnull()) & (data.HAD_CPOX<3)]["P_NUMVRC"])
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
data['P_NUMVRC'].unique()
```

Output:
```
array([nan,  1.,  0.,  2.,  3.])
```

```python
data['HAD_CPOX'].unique()
```

Output:
```
array([ 2,  1, 77, 99])
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
corr_chickenpox()
```

Output:
```
0.07044873460147986
```

```python

```